```python
# Initialize Otter
import otter
grader = otter.Notebook("hw3.ipynb")
```

```python
import numpy as np
```

# FILL IN YOUR NAME AND THE NAME OF YOUR PEER (IF ANY) BELOW

**Name**: \<replace this with your name\>

**Peer**: \<replace this with your peer's name\>

## Collaboration policy
Students are responsible for writing their own quizzes, assignments, projects, and exams. For homework assignments, students are welcome (and encouraged) to discuss problems with one peer, **but each student must write their own assignment writeup and code individually**. The peer must be listed at the top of the writeup for each assignment. *Note: Homework assignments are the primary learning mechanism for this course—you learn by doing. I strongly discourage the use of AI assistants in solving assignments. That said, I will treat AI assistants as peers: you are allowed to discuss problems with an AI assistant, but it is considered cheating to directly obtain an answer by querying the assistant. Please credit any AI assistant that you use.*

# Homework 3 — Motion planning (100 pts)

**Due:** Thursday, February 19th, 2026 at 11:59 pm

This homework builds on the material in the slides and Lynch Chapter 10 (on Motion Planning).

We will use Jupyter/Colab notebooks throughout the semester for writing code and generating assignment outputs.

**Throughout this homework, please use 2D numpy arrays to represent all matrices. For answers you must specify manually, the autograder will use a precision of 3 decimal places. We define a positive rotation as a counter-clockwise rotation.**


## 1) C-space obstacles (40 pts)

### 1.1) 2D polygons (20 pts)

In the case of 2D polygonal robots and obstacles, where the robot is only allowed translational motions, we can find the C-space obstacles by sliding the robot around the obstacle and keeping track of the origin of the robot as it touches the boundaries of the obstacle. The C-space obstacle is the polygon formed by connecting the places were the robot's origin was when the robot came in contact with the obstacle.


#### 1.1.1) 

Consider the following robot:

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q11_robot.png" width="30%">

and obstacle:

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q11_obstacle.png" width="40%">

The following are all representations of the C-space obstacle for different choices of the origin of the robot. Find the value of $(x,y)$ that corresponds to each of the C-space obstacles. In other words, determine what is the point on the robot where its origin is located, and then express the bottom-left corner of the robot relative to this origin. Assume that the x-axis points right, and the y-axis points up.

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q11_cspace_1.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q11_cspace_2.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q11_cspace_3.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q11_cspace_4.png" width="24%">



What is the value of $x,y$ that yields the first C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1111 = ...
y_1111 = ...
```

What is the value of $x,y$ that yields the second C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1112 = ...
y_1112 = ...
```

What is the value of $x,y$ that yields the third C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1113 = ...
y_1113 = ...
```

What is the value of $x,y$ that yields the fourth C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1114 = ...
y_1114 = ...
```

#### 1.1.2) 

Consider the following robot:

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q12_robot.png" width="30%">

and obstacle:

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q12_obstacle.png" width="30%">

The following are all representations of the C-space obstacle for different choices of the origin of the robot. Find the value of $(x,y)$ that corresponds to each of the C-space obstacles. In other words, determine what is the point on the robot where its origin is located, and then express the bottom-left corner of the robot relative to this origin. Assume that the x-axis points right, and the y-axis points up.

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q12_cspace_1.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q12_cspace_2.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q12_cspace_3.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q12_cspace_4.png" width="24%">



What is the value of $x,y$ that yields the first C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1121 = ...
y_1121 = ...
```

What is the value of $x,y$ that yields the second C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1122 = ...
y_1122 = ...
```

What is the value of $x,y$ that yields the third C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1123 = ...
y_1123 = ...
```

What is the value of $x,y$ that yields the fourth C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1124 = ...
y_1124 = ...
```

#### 1.1.3) 

Consider the following robot:

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q13_robot.png" width="30%">

and obstacle:

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q13_obstacle.png" width="40%">

The following are all representations of the C-space obstacle for different choices of the origin of the robot. Find the value of $(x,y)$ that corresponds to each of the C-space obstacles. In other words, determine what is the point on the robot where its origin is located, and then express the bottom-left corner of the robot relative to this origin. Assume that the x-axis points right, and the y-axis points up.

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q13_cspace_1.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q13_cspace_2.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q13_cspace_3.png" width="24%"> <img src="https://jorge-a-mendez.github.io/images/ese564s24/q13_cspace_4.png" width="24%">



What is the value of $x,y$ that yields the first C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1131 = ...
y_1131 = ...
```

What is the value of $x,y$ that yields the second C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1132 = ...
y_1132 = ...
```

What is the value of $x,y$ that yields the third C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1133 = ...
y_1133 = ...
```

What is the value of $x,y$ that yields the fourth C-space obstacle?

_Points:_ 1

```python
# Express your answer as two scalars
x_1134 = ...
y_1134 = ...
```

#### 1.1.4) 

Consider the following partial view of a robot's workspace with an obstacle: 

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q14_workspace.png" width="30%">

The configuration of the robot is given by the (x,y) coordinates of its bottom-left corner. Assume that the workspace is unbounded (i.e., it extends infinitely to the left, right, up, and down).


Give a configuration of the robot that is in C-free and for which $x>1$ and $y>2$. 

_Points:_ 2

```python
# Express your answer as two scalars
x_1141 = ...
y_1141 = ...
```

Give a configuration of the robot that is in C-obs. 

_Points:_ 2

```python
# Express your answer as two scalars
x_1142 = ...
y_1142 = ...
```

#### 1.1.5) 

Consider the following partial view of a robot's workspace with an obstacle: 

<img src="https://jorge-a-mendez.github.io/images/ese564s24/q15_workspace.png" width="30%">

The configuration of the robot is given by the (x,y) coordinates of its bottom-left corner. Assume that the workspace is unbounded (i.e., it extends infinitely to the left, right, up, and down).


Give a configuration of the robot that is in C-free and for which $x>1.8$ and $y>1.8$. 

_Points:_ 2

```python
# Express your answer as two scalars
x_1151 = ...
y_1151 = ...
```

Give a configuration of the robot that is in C-obs. 

_Points:_ 2

```python
# Express your answer as two scalars
x_1152 = ...
y_1152 = ...
```

### 1.2) 2-link planar arm (20 pts)

Consider the case of the 2-link planar arm with two revolute joints, as depicted below.

<img src="https://jorge-a-mendez.github.io/images/ese564s24/twolink_workspace.png" width="30%">

The base of link 1 as at the origin of the workspace, $(0,0)$. The links of the arm both have length $L_1=L_2=3$. There are three obstacles in the workspace, with vertices annotated in the figure. The configuration is given by the angles of the joints, $q_1$ and $q_2$.

#### 1.2.1) 
Which of the following represents the configuration space of the robot? Black pixels represent free space, and white pixels represent obstacles.

*Hint: consider writing your own piece of code that visualizes the robot in your choice of configurations.*

<img src="https://jorge-a-mendez.github.io/images/ese564s24/twolink_cspace_option1.png" width="24%">
<img src="https://jorge-a-mendez.github.io/images/ese564s24/twolink_cspace_option2.png" width="24%">
<img src="https://jorge-a-mendez.github.io/images/ese564s24/twolink_cspace_option3.png" width="24%">
<img src="https://jorge-a-mendez.github.io/images/ese564s24/twolink_cspace_option4.png" width="24%">

_Points:_ 3

```python
# Express your answer as a lower-case Python string (e.g., "a")
ans_121 = ...
```

#### 1.2.2) 

How many connected components does this C-space have?

_Points:_ 3

```python
# Express your answer as a scalar integer
ans_122 = ...
```

#### 1.2.3) 

Give an example of a configuration $q_1, q_2$ that is in C-free.

_Points:_ 3

```python
# Express your answer as two scalars in [0, 2pi]
q1_123 = ...
q2_123 = ...
```

#### 1.2.4) 

Give an example of a configuration $q_1, q_2$ that is in C-obs.

_Points:_ 3

```python
# Express your answer as two scalars in [0, 2pi]
q1_124 = ...
q2_124 = ...
```

#### 1.2.5) 

Give an example of a pair of configurations $q_\text{start}$, $q_\text{goal}$ in C-free such that there is no collision-free path from start to goal.

_Points:_ 4

```python
# Express your answer as four scalars in [0, 2pi]
q1start_125 = ...
q2start_125 = ...
q1goal_125 = ...
q2goal_125 = ...
```

#### 1.2.6) 

Give an example of a pair of configurations $q_\text{start}$, $q_\text{goal}$ in C-free such that there is some collision-free path from start to goal.

_Points:_ 4

```python
# Express your answer as four scalars in [0, 2pi]
q1start_126 = ...
q2start_126 = ...
q1goal_126 = ...
q2goal_126 = ...
```

## 2) RRT implementation (60 pts)

We are now going to write an implementation of RRT for the Panda robot. Run the code cell below to initialize the relevant variables and visualize the environment with obstacles.

The following three functions defined below may be particularly useful for your implementation and debugging: `set_joint_positions` and `get_joint_positions` are self-explanatory, hopefully. `get_home_joint_positions` returns values of the joint positions where the Panda is in a "default" configuration. `get_goal_joint_positions` returns values of the joint positions where the Panda is no longer between the obstacles. You can use these to debug your overall RRT implementation.

**Do not modify any paths in the cell below, as this code cell must run exactly as is on the autograder, which assumes the paths are fixed to the values set below.**

**If you work on Google Colab, you will likely need to set up virtual rendering to be able to visualize things. If you do, please remove any cells that set up virtual rendering before submitting your assignment, as virtual rendering is not allowed on Gradescope and it will likely cause the autograder to fail on all subsequent questions.**

```python
import mujoco
import mediapy as media
import os
import glfw

def has_valid_renderer():
    """Checks if OpenGL can be initialized successfully."""
    # Try initializing GLFW
    if not glfw.init():
        print("GLFW initialization failed. No valid renderer.")
        return False
    
    # Try creating a hidden window to test OpenGL context creation
    window = glfw.create_window(640, 480, "Renderer Test", None, None)
    if not window:
        print("GLFW could not create an OpenGL context. No valid renderer.")
        glfw.terminate()
        return False
    
    glfw.make_context_current(window)
    renderer_available = True  # If we reach this point, rendering should work

    # Cleanup
    glfw.destroy_window(window)
    glfw.terminate()
    return renderer_available

def render_state(model, data):
    if not has_valid_renderer():
        print("Error rendering the scene. Make sure you have a valid OpenGL installation.")
        return

    mujoco.mj_forward(model, data)
    with mujoco.Renderer(model, height=480, width=640) as renderer:
        renderer.update_scene(data, "front_view")  # Ensure scene is updated with final simulation state
        media.show_image(renderer.render())

def set_joint_positions(model, data, joints):
    joint_ids = [mujoco.mj_name2id(model, mujoco.mjtObj.mjOBJ_JOINT, f"panda_joint{i+1}") for i in range(7)]
    joint_qpos_idx = model.jnt_qposadr[joint_ids]
    data.qpos[joint_qpos_idx] = joints
    mujoco.mj_forward(model, data)

def get_joint_positions(model, data):
    joint_ids = [mujoco.mj_name2id(model, mujoco.mjtObj.mjOBJ_JOINT, f"panda_joint{i+1}") for i in range(7)]
    joint_qpos_idx = model.jnt_qposadr[joint_ids]
    return data.qpos[joint_qpos_idx]

def get_home_joint_positions(model, data):
    return [0, -np.pi/3, 0, -5*np.pi/6, 0, np.pi/2, 0]

def get_goal_joint_positions(model, data):
    return [-np.pi/2, np.pi/6, 0, -np.pi/3, np.pi/2, np.pi/2, 0]
    
# Define the MJCF path and mesh directory
mjcf_path = 'assets/Panda/panda.mjcf'
mesh_dir = 'assets/Panda/meshes/collision/'

# Load mesh files into a dictionary
assets = {}
for file_name in os.listdir(mesh_dir):
    file_path = os.path.join(mesh_dir, file_name)
    with open(file_path, 'rb') as f:
        assets[f'meshes/collision/{file_name}'] = f.read()

# Load the MuJoCo model
model = mujoco.MjModel.from_xml_path(mjcf_path, assets=assets)
data = mujoco.MjData(model)
gripper_joint_ids = [
    mujoco.mj_name2id(model, mujoco.mjtObj.mjOBJ_JOINT, "panda_finger_joint1"),
    mujoco.mj_name2id(model, mujoco.mjtObj.mjOBJ_JOINT, "panda_finger_joint2")
]
gripper_joint_qpos_idx = model.jnt_qposadr[gripper_joint_ids]
data.qpos[gripper_joint_qpos_idx] = 0.04

# Call mj_forward to initialize the state
mujoco.mj_forward(model, data)

# Set the joint positions
set_joint_positions(model, data, get_home_joint_positions(model, data))

# Render the scene
render_state(model, data)
```

##### 2.1)

Implement the `sample_point_21` function. With probability $1-\epsilon$, the function should return a joint configuration sampled uniformly at random from the valid joint ranges. With probability $\epsilon$, the function should return the goal configuration.


_Points:_ 10

```python
def sample_point_21(model, q_goal, epsilon=0.1):
    """ 
    model: MuJoCo model object
    q_goal: goal joint positions, a numpy array of shape (7,)
    epsilon: probability of sampling the goal configuration
    
    Return: a numpy array of shape (7,) representing the sampled joint positions
    """
    ...
```

```python
grader.check("q2.1")
```

### 2.2) 

Let's now implement the `get_nearest_22` function. Because all Panda joints have joint limits, there are no wrap-arounds and we can get away with using Euclidean distance as our distance metric. 

_Points:_ 10

```python
def get_nearest_22(nodes, q):
    """ 
    nodes: a numpy array of shape (N, 7) representing the nodes in the tree
    q: a numpy array of shape (7,) representing the query point
    Return: a numpy array of shape (7,) representing the nearest node in the tree
    """
    ...
```

```python
grader.check("q2.2")
```

#### 2.3) 
The next step is to write the `extend_23` function, which findes `q_new`, a node on the line between `q_sample` and `q_near` that is at a maximum distance of `step_size` from `q_near` (again, using Euclidean distance in the C-space). If `q_sample` is closer to `q_near` than `step_size`, then `q_new=q_sample`.

_Points:_ 10

```python
def extend_23(q_sample, q_near, step_size):
    """ 
    q_sample: a numpy array of shape (7,) representing the sampled joint positions
    q_near: a numpy array of shape (7,) representing the nearest node in the tree
    step_size: a scalar representing the step size
    Return: a numpy array of shape (7,) representing the extended joint positions
    """
    ...
```

```python
grader.check("q2.3")
```

#### 2.4)

Before moving on to write the main loop of RRT, we need to write a collision-checking function. The function takes as input two configurations, q1 and q2, and checks whether the path between them is collision-free. The way we will do this is to interpolate between the configurations with a fixed step size.

To check whether a particular configuration is in collision, we will leverage the MuJoCo simulator. In particular, the following will be useful:
- `mujoco.mj_collision(model, data)` runs collision checking internally in MuJoCo. 
- `data.contact` is set after running `mj_collision`. Look at the documentation to see how to use this information to determine whether there are collisions.

*Hint1: this function will require you to set the joint positions multiple times. Make sure that, at the end of the function, your robot is back to its original position (not q1, but the position of the robot in the simulator before entering collision checking).*

*Hint2: it is possible that the two configurations are closer than the `max_step_size`. Make sure to handle this case appropriately.*

*Hint3: if you use `np_array += other_array`, you will modify `np_array` in-place.*

_Points:_ 10

```python
def is_collision_free_24(model, data, q1, q2, max_step_size=0.1):
    """ 
    model: MuJoCo model object
    data: MuJoCo data object
    q1: a numpy array of shape (7,) representing the start joint positions
    q2: a numpy array of shape (7,) representing the end joint positions
    max_step_size: a scalar representing the maximum step size for collision checking
    
    Return: a boolean indicating if the path between q1 and q2 is collision-free
    """
    ...
```

```python
grader.check("q2.4")
```

#### 2.5)

Now we're ready to write our RRT main loop. The algorithm loops over: sampling a point, finding the nearest node on the tree, finding the candidate extension node, checking for collisions, and adding to the tree.

The function will take as arguments a threshold `threshold` which tells us how close (in Euclidean distance in the C-space) we must be from the goal configuration to consider that the goal has been reached, and a `step_size` to use for the `extend` function. For the purposes of this HW, feel free to use the same `step_size` for collision checking (which effectively just checks whether `q_new` is in collision).

Recall that RRT keeps track of the parent of each node. You may use a `parent` dictionary that contains each node as a key, and the parent of each key node as its value. (e.g., `{child_node: parent_node}`). You will use this dictionary to retrace the plan by finding the parent of the current node recursively, starting from the goal.

*Hint 1: Numpy arrays cannot be used as keys in a dictionary. You may need to convert back and forth between tuples and arrays. Make sure to respect the typing expected by the functions defined so far.*

*Hint 2: You can visualize your plans by calling `render_state`, `time.sleep(0.1)` and `IPython.display.clear_output`*


_Points:_ 10

```python
def retrace_plan_25(q_goal, tree):
    """ 
    q_goal: a numpy array of shape (7,) representing the goal joint positions
    tree: a dictionary mapping each node (as a tuple) to its parent node (as a tuple)
    Return: a list of numpy arrays representing the sequence of joint positions from q_init to q_goal.
    """
    ...

def panda_rrt_25(q_init, q_goal, model, data, threshold=0.01, step_size=0.1):
    """ 
    q_init: a numpy array of shape (7,) representing the initial joint positions
    q_goal: a numpy array of shape (7,) representing the goal joint positions
    model: MuJoCo model object
    data: MuJoCo data object
    threshold: a scalar representing the distance threshold to the goal
    step_size: a scalar representing the step size for extending the tree

    Return: a list of numpy arrays representing the sequence of joint positions from q_init to q_goal.    
        The first element should be q_init and the final element should be within threshold of q_goal
    """
    ...
```

```python
grader.check("q2.5")
```

#### 2.6)

When using a standard implementation of RRT, it is common practice to apply some smoothing mechanism on the plan obtained from RRT. In this question, we will explore a simple smoothing algorithm.

The pseudocode is the following:
```
for i in 1 ... max_iters
    i, j <- sample non-consecutive indices i < j + 1 that are neither the first nor the final in the plan
    if q[i] -> q[j] can be connected by a straight-line collision-free path, remove all intermediate steps from the plan
```

Implement this algorithm as `smooth_path_26` below.

_Points:_ 10

```python
def smooth_path_26(path, model, data, max_iters=1000):
    """
    path: a list of numpy arrays representing the sequence of joint positions from q_init to q_goal
    model: MuJoCo model object
    data: MuJoCo data object
    max_iters: maximum number of iterations for smoothing

    Return: a list of numpy arrays representing the smoothed path 
    """
    ...
```

```python
grader.check("q2.6")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Fill out the answers to all questions and submit your file hw3.ipynb to the HW3 assignment on Gradescope. You are free to resubmit as many times as you wish.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1.1.1.1 (points: 1)


Tests for q1.1.1.2 (points: 1)


Tests for q1.1.1.3 (points: 1)


Tests for q1.1.1.4 (points: 1)


Tests for q1.1.2.1 (points: 1)


Tests for q1.1.2.2 (points: 1)


Tests for q1.1.2.3 (points: 1)


Tests for q1.1.2.4 (points: 1)


Tests for q1.1.3.1 (points: 1)


Tests for q1.1.3.2 (points: 1)


Tests for q1.1.3.3 (points: 1)


Tests for q1.1.3.4 (points: 1)


Tests for q1.1.4.1 (points: 2)


Tests for q1.1.4.2 (points: 2)


Tests for q1.1.5.1 (points: 2)


Tests for q1.1.5.2 (points: 2)


Tests for q1.2.1 (points: 3)


Tests for q1.2.2 (points: 3)


Tests for q1.2.3 (points: 3)


Tests for q1.2.4 (points: 3)


Tests for q1.2.5 (points: 4)


Tests for q1.2.6 (points: 4)


Tests for q2.1 (points: 10)
```
>>> q_goal = get_goal_joint_positions(model, data)
>>> q_sample = sample_point_21(model, q_goal)
>>> isinstance(q_sample, np.ndarray) and q_sample.shape == (7,)
True
```
failure_message: Did not return a numpy array of shape (7,)
```
>>> q_goal = get_goal_joint_positions(model, data)
>>> eps = 1.0
>>> all_goal_list = []
>>> for _ in range(10):
...     q_sample = sample_point_21(model, q_goal, epsilon=eps)
...     all_goal_list.append(np.allclose(q_sample, q_goal))
>>> all(all_goal_list)
True
```
failure_message: With eps = 1.0, did not always return the goal configuration
```
>>> q_goal = get_goal_joint_positions(model, data)
>>> eps = 0.0
>>> all_nogoal_list = []
>>> for _ in range(10):
...     q_sample = sample_point_21(model, q_goal, epsilon=eps)
...     all_nogoal_list.append(not np.allclose(q_sample, q_goal))
>>> all(all_nogoal_list)
True
```
failure_message: With eps = 0.0, did not always return a non-goal configuration


Tests for q2.2 (points: 10)
```
>>> nodes = np.array([[0, 0, 0, 0, 0, 0, 0], [1, 1, 1, 1, 1, 1, 1], [2, 2, 2, 2, 2, 2, 2]])
>>> q = np.array([1.1, 1.1, 1.1, 1.1, 1.1, 1.1, 1.1])
>>> nearest_node = get_nearest_22(nodes, q)
>>> isinstance(nearest_node, np.ndarray) and nearest_node.shape == (7,)
True
```
failure_message: Did not return a numpy array of shape (7,)
```
>>> nodes = np.array([[0, 0, 0, 0, 0, 0, 0], [1, 1, 1, 1, 1, 1, 1], [2, 2, 2, 2, 2, 2, 2]])
>>> q = np.array([1.1, 1.1, 1.1, 1.1, 1.1, 1.1, 1.1])
>>> nearest_node = get_nearest_22(nodes, q)
>>> np.allclose(nearest_node, np.array([1, 1, 1, 1, 1, 1, 1]))
True
```
failure_message: Did not return the nearest node


Tests for q2.3 (points: 10)
```
>>> q_sample = np.array([1, 1, 1, 1, 1, 1, 1])
>>> q_near = np.array([0, 0, 0, 0, 0, 0, 0])
>>> step_size = 0.5
>>> q_new = extend_23(q_sample, q_near, step_size)
>>> isinstance(q_new, np.ndarray) and q_new.shape == (7,)
True
```
failure_message: Did not return a numpy array of shape (7,)
```
>>> q_sample = np.array([1, 1, 1, 1, 1, 1, 1])
>>> q_near = np.array([0, 0, 0, 0, 0, 0, 0])
>>> step_size = 0.5
>>> q_new = extend_23(q_sample, q_near, step_size)
>>> np.allclose(q_new, np.array([0.5 / np.sqrt(7)] * 7))
True
```
failure_message: Did not return the correct array when distance > step_size
```
>>> q_sample = np.array([1, 1, 1, 1, 1, 1, 1])
>>> q_near = np.array([0, 0, 0, 0, 0, 0, 0])
>>> step_size = 3
>>> q_new = extend_23(q_sample, q_near, step_size)
>>> np.allclose(q_new, q_sample)
True
```
failure_message: Did not return the correct array when distance <= step_size


Tests for q2.4 (points: 10)
```
>>> q1 = np.array(get_home_joint_positions(model, data))
>>> q2 = np.array(get_goal_joint_positions(model, data))
>>> collision_free = is_collision_free_24(model, data, q1, q2)
>>> isinstance(collision_free, bool)
True
```
failure_message: Did not return a boolean scalar
```
>>> q1 = get_home_joint_positions(model, data)
>>> q2 = np.array(get_goal_joint_positions(model, data))
>>> set_joint_positions(model, data, q2)
>>> _ = is_collision_free_24(model, data, q1, q2)
>>> final_q = np.array(get_joint_positions(model, data))
>>> np.allclose(final_q, q2)
True
```
failure_message: Did not return the robot to its original configuration
```
>>> joint_ids = [mujoco.mj_name2id(model, mujoco.mjtObj.mjOBJ_JOINT, f'panda_joint{i + 1}') for i in range(7)]
>>> joint_ranges = np.array([model.jnt_range[joint_id] for joint_id in joint_ids])
>>> q_base = np.random.uniform(joint_ranges[:, 0], joint_ranges[:, 1], size=7)
>>> set_joint_positions(model, data, q_base)
>>> q1 = np.random.uniform(joint_ranges[:, 0], joint_ranges[:, 1], size=7)
>>> q2 = np.random.uniform(joint_ranges[:, 0], joint_ranges[:, 1], size=7)
>>> max_step_size = 0.1
>>> _ = is_collision_free_24(model, data, q1, q2, max_step_size)
>>> np.allclose(get_joint_positions(model, data), q_base)
True
```


Tests for q2.5 (points: 10)
```
>>> q_init = np.array(get_home_joint_positions(model, data))
>>> q_goal = np.array(get_goal_joint_positions(model, data))
>>> plan = panda_rrt_25(q_init, q_goal, model, data)
>>> isinstance(plan, list) and all((isinstance(q, np.ndarray) and q.shape == (7,) for q in plan))
True
```
failure_message: Did not return a list of numpy arrays of shape (7,)
```
>>> q_init = np.array(get_home_joint_positions(model, data))
>>> q_goal = np.array(get_goal_joint_positions(model, data))
>>> plan = panda_rrt_25(q_init, q_goal, model, data)
>>> np.allclose(plan[0], q_init)
True
```
failure_message: The first element of the plan is not q_init
```
>>> q_init = np.array(get_home_joint_positions(model, data))
>>> q_goal = np.array(get_goal_joint_positions(model, data))
>>> plan = panda_rrt_25(q_init, q_goal, model, data)
>>> np.linalg.norm(plan[-1] - q_goal).item() < 0.01
True
```
failure_message: The final element of the plan is not q_goal


Tests for q2.6 (points: 10)
```
>>> q_init = np.array(get_home_joint_positions(model, data))
>>> q_goal = np.array(get_goal_joint_positions(model, data))
>>> plan = panda_rrt_25(q_init, q_goal, model, data)
>>> smoothed_plan = smooth_path_26(plan, model, data)
>>> isinstance(smoothed_plan, list) and all((isinstance(q, np.ndarray) and q.shape == (7,) for q in smoothed_plan))
True
```
failure_message: Did not return a list of numpy arrays of shape (7,)
```
>>> q_init = np.array(get_home_joint_positions(model, data))
>>> q_goal = np.array(get_goal_joint_positions(model, data))
>>> plan = panda_rrt_25(q_init, q_goal, model, data)
>>> smoothed_plan = smooth_path_26(plan, model, data)
>>> np.allclose(smoothed_plan[0], q_init)
True
```
failure_message: The first element of the plan is not q_init
```
>>> q_init = np.array(get_home_joint_positions(model, data))
>>> q_goal = np.array(get_goal_joint_positions(model, data))
>>> plan = panda_rrt_25(q_init, q_goal, model, data)
>>> smoothed_plan = smooth_path_26(plan, model, data)
>>> np.linalg.norm(smoothed_plan[-1] - q_goal).item() < 0.01
True
```
failure_message: The final element of the plan is not q_goal
